### Codio Activity 22.8: Hyperparameter Tuning with `keras`

**Expected Time = 60 minutes**

**Total Points = 40**

This activity focuses on using hyperparameter tuning with the keras library.  There are two ways to perform a grid search with keras, and you will implement both. While `keras_tuner` was discussed in the lectures, here you will use the scikit-learn wrapper for keras where you can grid search the parameters using `GridSearchCV`.  You will implement this with the `KerasClassifier` to build some basic models on the wine dataset.  

#### Index

- [Problem 1](#-Problem-1)
- [Problem 2](#-Problem-2)
- [Problem 3](#-Problem-3)
- [Problem 4](#-Problem-4)

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd

from tensorflow.keras.wrappers.scikit_learn import KerasClassifier
from tensorflow.keras.layers import Dense
from tensorflow.keras.models import Sequential
from tensorflow.keras.utils import to_categorical
import tensorflow as tf

from sklearn.datasets import load_wine
from sklearn.preprocessing import StandardScaler
from sklearn.model_selection import train_test_split, GridSearchCV

import warnings
```

### The Data

Below, the wine dataset is loaded, split, and scaled.  

```python
wine = load_wine(as_frame=True)
```

```python
wine.frame.head()
```

```python
X = wine.data
y = to_categorical(wine.target)
```

```python
X_scaled = StandardScaler().fit_transform(X)
```

[Back to top](#-Index)

### Problem 1

#### The Build Function

**10 Points**

To use the `KerasClassifier` you first need to write a function that creates a keras model and takes in arguments for the parameters you wish to search. Below is a very basic example.  Note that rather than 

```python
def build_function(neurons = 1):
    model = Sequential()
    model.add(Dense(neurons, activation = 'relu'))
    model.add(Dense(3, activation = 'softmax'))
    model.compile(...)
    return model
```

Complete the `build_function` below with appropriate compilation arguments.



<!-- grade_id: cell-a2bce77d17112fd1, points: 0, grade: false, solution: true -->
```python
### GRADED
tf.random.set_seed(42)
def build_function(neurons = 1):
    model = ''
    #model.compile()
    return model

    
### BEGIN SOLUTION
tf.random.set_seed(42)
def build_function(neurons = 1):
    model = Sequential()
    model.add(Dense(neurons, activation = 'relu'))
    model.add(Dense(3, activation = 'softmax'))
    model.compile(loss = 'categorical_crossentropy', metrics = ['accuracy'])
    return model
### END SOLUTION

### ANSWER CHECK
build_function
```

<!-- grade_id: cell-817b7e106ad4c353, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tf.random.set_seed(42)
def build_function_(neurons = 1):
    model = Sequential()
    model.add(Dense(neurons, activation = 'relu'))
    model.add(Dense(3, activation = 'softmax'))
    model.compile(loss = 'categorical_crossentropy', metrics = ['accuracy'])
    return model
ans = build_function_(neurons = 100)
stu = build_function(neurons = 100)
#
#
#
assert ans.layers[0].units == stu.layers[0].units
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 2

#### Creating the `KerasClassifier`

**10 Points**

Now, use the `build_function` to instantiate the `KerasClassifier` as `model` with `epochs = 10`. Fit the model on the data `X_scaled` and `y`, and assign the accuracy as `score_1` below.  Note that this will use 1 node as defaulted in the `build_function`.

<!-- grade_id: cell-0275116e4f507015, points: 0, grade: false, solution: true -->
```python
### GRADED
tf.random.set_seed(42)
model = ''
score_1 = ''

    
### BEGIN SOLUTION
tf.random.set_seed(42)
model = KerasClassifier(build_fn=build_function, 
                        epochs = 10, 
                       verbose = 0)
history = model.fit(X_scaled, y)
model.score(X_scaled, y)
### END SOLUTION

### ANSWER CHECK
print(model.score(X_scaled, y))
```

<!-- grade_id: cell-7696f2d368bb0be4, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tf.random.set_seed(42)
model_ = KerasClassifier(build_fn=build_function_, 
                        epochs = 10, 
                       verbose = 0)
history = model_.fit(X_scaled, y)
ans = model_.score(X_scaled, y)
stu = model.score(X_scaled, y)
#
#
#
assert ans == stu
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 3

#### Performing the Grid Search

**10 Points**

Now, to perform a grid search you just need to create a dictionary with the hyperparameter named `neurons = [10, 50, 100]`.  Use `GridSearchCV` with the model and appropriate hyperparameter dictionary to fit the data.  

```python
ann_grid = GridSearchCV(model, param_grid = params)
```

<!-- grade_id: cell-8c00e8a1dee417d8, points: 0, grade: false, solution: true -->
```python
### GRADED
tf.random.set_seed(42)
params = ''
ann_grid = ''

    
### BEGIN SOLUTION
tf.random.set_seed(42)
params = {'neurons': [10, 50, 100]}
ann_grid = GridSearchCV(model, param_grid=params)
### END SOLUTION

### ANSWER CHECK
ann_grid
```

<!-- grade_id: cell-addbffa1795cc12f, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tf.random.set_seed(42)
params_ = {'neurons': [10, 50, 100]}
ann_grid_ = GridSearchCV(model_, param_grid=params_)
#
#
#
assert ann_grid.get_params()['param_grid'] == ann_grid_.get_params()['param_grid']
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 4

#### Fit and Evaluate the model

**10 Points**

Now, fit and evaluate the model.  Assign your score to `grid_score` below. 

<!-- grade_id: cell-f8e6edb517c71335, points: 0, grade: false, solution: true -->
```python
### GRADED
tf.random.set_seed(42)

grid_score = ''

    
### BEGIN SOLUTION
tf.random.set_seed(42)
params = {'neurons': [10, 50, 100]}
ann_grid = GridSearchCV(model, param_grid=params)
grid_score = ann_grid.fit(X_scaled, y).score(X_scaled, y)
### END SOLUTION

### ANSWER CHECK
print(grid_score)
```

<!-- grade_id: cell-fbf8f8e25550261d, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tf.random.set_seed(42)
params_ = {'neurons': [10, 50, 100]}
ann_grid_ = GridSearchCV(model_, param_grid=params_)
grid_score_ = ann_grid_.fit(X_scaled, y).score(X_scaled, y)
#
#
#
assert grid_score == grid_score_
### END HIDDEN TESTS
```

Because of the grading enviornment, a more exhaustive search over additional parameters is not an option.  To extend the work here should be straightforward enough, and this is a nice solution to grid searching the hyperparameters of a model. 