### Codio Activity 23.4: Fine Tuning a Pre-trained Network

**Expected Time = 60 minutes**

**Total Points = 30**

In addition to the use of a pre-trained network to extract features from a different dataset, the weights can be adjusted or **fine-tuned** as a last step to squeeze additional performance from the network.  To do so, you will again use the `EfficientNetV2B0` network on the `cifar10` data from `keras`.  This time you are encouraged to use the functional API syntax to construct your network.  

For a second example, consult the `keras` documentation example [here](https://keras.io/guides/transfer_learning/). 

#### Index

- [Problem 1](#-Problem-1)
- [Problem 2](#-Problem-2)
- [Problem 3](#-Problem-3)

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt


from tensorflow.keras.applications.efficientnet_v2 import EfficientNetV2B0
from tensorflow.keras.layers import Dense, Flatten
from tensorflow.keras.models import Sequential
from tensorflow.keras.utils import to_categorical
import tensorflow as tf
from tensorflow.keras.datasets import cifar10
from tensorflow.keras import Model, Input
```

Output:
```
2022-08-03 20:50:38.514197: W tensorflow/stream_executor/platform/default/dso_loader.cc:64] Could not load dynamic library 'libcudart.so.11.0'; dlerror: libcudart.so.11.0: cannot open shared object file: No such file or directory
2022-08-03 20:50:38.514237: I tensorflow/stream_executor/cuda/cudart_stub.cc:29] Ignore above cudart dlerror if you do not have a GPU set up on your machine.
```

```python
(X_train, y_train), (X_test, y_test) = cifar10.load_data()
sub_indices = pd.DataFrame(y_train, columns= ["labels"]).groupby('labels', group_keys=False).apply(lambda x: x.sample(frac=0.3, random_state=123)).index
X_train = X_train[sub_indices]
y_train = y_train[sub_indices]

Y_train = to_categorical(y_train)
Y_test = to_categorical(y_test)

```

[Back to top](#-Index)

### Problem 1

#### Training the Network to Convergence

**10 Points**

Below, use the `EfficientNetV2B0` as `base_model` with appropriate input shape and the functional API to create a model using the `base_model`, flatten the results, pass through a `Dense` layer of 100 units and an appropriate output layer.  Name the `Input` inputs and the output layer `outputs`.  Compile the model and fit it using 20 epochs, assigning the history as `history`.  

**Make sure the `base_model` weights are not trainable.**

<!-- grade_id: cell-e92e20d0c85d3026, points: 0, grade: false, solution: true -->
```python
### GRADED
base_model = ''
inputs = ''
x = ''
x = ''
x = ''
output = ''
model = ''
#be sure to compile

bottom_model = ''
    
### BEGIN SOLUTION
base_model = EfficientNetV2B0(input_shape = (32, 32, 3), include_top=False)
inputs = Input(shape = (32, 32, 3))
x = base_model(inputs)
x = Flatten()(x)
x = Dense(100, activation = 'relu')(x)
output = Dense(10, activation = 'sigmoid')(x)
base_model.trainable = False
model = Model(inputs, output)
model.compile(loss = 'categorical_crossentropy', metrics = ['accuracy'])
bottom_model = model.fit(X_train, Y_train, validation_data=(X_test, Y_test), 
                    epochs = 20)
### END SOLUTION

### ANSWER CHECK
print(model.summary())
```

Output:
```
2022-08-03 20:50:44.170421: W tensorflow/stream_executor/platform/default/dso_loader.cc:64] Could not load dynamic library 'libcuda.so.1'; dlerror: libcuda.so.1: cannot open shared object file: No such file or directory
2022-08-03 20:50:44.170485: W tensorflow/stream_executor/cuda/cuda_driver.cc:269] failed call to cuInit: UNKNOWN ERROR (303)
2022-08-03 20:50:44.170512: I tensorflow/stream_executor/cuda/cuda_diagnostics.cc:156] kernel driver does not appear to be running on this host (victorbravo-kineticenjoy): /proc/driver/nvidia/version does not exist
2022-08-03 20:50:44.171293: I tensorflow/core/platform/cpu_feature_guard.cc:193] This TensorFlow binary is optimized with oneAPI Deep Neural Network Library (oneDNN) to use the following CPU instructions in performance-critical operations:  AVX2 AVX512F FMA
To enable them in other operations, rebuild TensorFlow with the appropriate compiler flags.
Epoch 1/20
469/469 [==============================] - 27s 46ms/step - loss: 1.5263 - accur
```

<!-- grade_id: cell-c7a974852163092b, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
base_model_ = EfficientNetV2B0(input_shape = (32, 32, 3), include_top=False)
inputs_ = Input(shape = (32, 32, 3))
x_ = base_model(inputs_)
x_ = Flatten()(x_)
x_ = Dense(100, activation = 'relu')(x_)
output_ = Dense(10, activation = 'sigmoid')(x_)
model_ = Model(inputs_, output_)
model_.compile(loss = 'categorical_crossentropy', metrics = ['accuracy'])
bottom_model_ = model_.fit(X_train, Y_train, validation_data=(X_test, Y_test), 
                    epochs = 20)
#
#
#
assert type(base_model) == type(base_model_)
assert type(bottom_model) == type(bottom_model_)
### END HIDDEN TESTS
```

Output:
```
Epoch 1/20
469/469 [==============================] - 31s 49ms/step - loss: 1.5299 - accuracy: 0.4687 - val_loss: 1.3571 - val_accuracy: 0.5157
Epoch 2/20
469/469 [==============================] - 21s 44ms/step - loss: 1.2544 - accuracy: 0.5585 - val_loss: 1.2741 - val_accuracy: 0.5442
Epoch 3/20
469/469 [==============================] - 20s 43ms/step - loss: 1.1771 - accuracy: 0.5859 - val_loss: 1.1896 - val_accuracy: 0.5814
Epoch 4/20
469/469 [==============================] - 22s 48ms/step - loss: 1.1329 - accuracy: 0.5981 - val_loss: 1.1526 - val_accuracy: 0.5956
Epoch 5/20
469/469 [==============================] - 21s 45ms/step - loss: 1.0908 - accuracy: 0.6169 - val_loss: 1.1891 - val_accuracy: 0.5824
Epoch 6/20
469/469 [==============================] - 20s 42ms/step - loss: 1.0667 - accuracy: 0.6225 - val_loss: 1.1429 - val_accuracy: 0.6024
Epoch 7/20
469/469 [==============================] - 21s 46ms/step - loss: 1.0369 - accuracy: 0.6365 - val_loss: 1.1852 - val_accuracy:
```

[Back to top](#-Index)

### Problem 2

#### Setting to Trainable

**10 Points**

After fitting the model above, you can set the layers to trainable to update the weights.  Below, set the final five layers to trainable in the `base_model` as demonstrated in the lectures. 

```python
base_model.trainable #are the base model weights set to trainable?
```

Output:
```
False
```

<!-- grade_id: cell-17e4f0a4de6dc3b8, points: 0, grade: false, solution: true -->
```python
### GRADED
base_model.trainable = True
for layer in '':
    #make trainable
    pass
    
### BEGIN SOLUTION
base_model.trainable = True
for layer in base_model.layers[:-5]:
    layer.trainable = False
### END SOLUTION

### ANSWER CHECK
for i, layer in enumerate(base_model.layers[-10:]):
    print(f'Layer is trainable: {layer.trainable}')
```

Output:
```
Layer is trainable: False
Layer is trainable: False
Layer is trainable: False
Layer is trainable: False
Layer is trainable: False
Layer is trainable: True
Layer is trainable: True
Layer is trainable: True
Layer is trainable: True
Layer is trainable: True
```

<!-- grade_id: cell-279935493600bece, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
base_model_.trainable = True
for layer in base_model_.layers[:-5]:
    layer.trainable = False
#
#
#
for i, layer in enumerate(base_model.layers[-5:]):
    assert layer.trainable == True
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 3

#### Refitting the network

**10 Points**

Finally, specify the `model` compilation and fit this again for 10 epochs, assigning the resulting history to `fine_tuned_history` below. 

<!-- grade_id: cell-75f43fbafe9c0840, points: 0, grade: false, solution: true -->
```python
### GRADED
fine_tuned_history = ''
    
### BEGIN SOLUTION
model.compile(loss = 'categorical_crossentropy', metrics = ['accuracy'])
fine_tuned_history = model.fit(X_train, Y_train, validation_data = (X_test, Y_test), epochs = 10)
### END SOLUTION

### ANSWER CHECK
fine_tuned_history.history['accuracy'][-1]
```

Output:
```
Epoch 1/10
469/469 [==============================] - 39s 57ms/step - loss: 3.2024 - accuracy: 0.4689 - val_loss: 1.5707 - val_accuracy: 0.5584
Epoch 2/10
469/469 [==============================] - 20s 42ms/step - loss: 1.1948 - accuracy: 0.6214 - val_loss: 1.3980 - val_accuracy: 0.5848
Epoch 3/10
469/469 [==============================] - 20s 43ms/step - loss: 0.9943 - accuracy: 0.6783 - val_loss: 1.3801 - val_accuracy: 0.6092
Epoch 4/10
469/469 [==============================] - 20s 43ms/step - loss: 0.9086 - accuracy: 0.7088 - val_loss: 1.4395 - val_accuracy: 0.6109
Epoch 5/10
469/469 [==============================] - 20s 42ms/step - loss: 0.8263 - accuracy: 0.7309 - val_loss: 1.4466 - val_accuracy: 0.6098
Epoch 6/10
469/469 [==============================] - 21s 44ms/step - loss: 0.7448 - accuracy: 0.7548 - val_loss: 1.4951 - val_accuracy: 0.6130
Epoch 7/10
469/469 [==============================] - 22s 46ms/step - loss: 0.7017 - accuracy: 0.7729 - val_loss: 1.5713 - val_accuracy:
```

<!-- grade_id: cell-c60ddea5358927c2, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
model_.compile(loss = 'categorical_crossentropy', metrics = ['accuracy'])
fine_tuned_history_ = model_.fit(X_train, Y_train, validation_data = (X_test, Y_test), epochs = 10)
#
#
#
assert len(fine_tuned_history.history['accuracy']) == len(fine_tuned_history_.history['accuracy'])
### END HIDDEN TESTS
```

Output:
```
Epoch 1/10
469/469 [==============================] - 31s 52ms/step - loss: 1.0456 - accuracy: 0.6992 - val_loss: 1.6868 - val_accuracy: 0.6022
Epoch 2/10
469/469 [==============================] - 21s 45ms/step - loss: 0.7396 - accuracy: 0.7624 - val_loss: 1.7184 - val_accuracy: 0.6112
Epoch 3/10
469/469 [==============================] - 20s 42ms/step - loss: 0.6652 - accuracy: 0.7913 - val_loss: 1.7860 - val_accuracy: 0.6119
Epoch 4/10
469/469 [==============================] - 20s 43ms/step - loss: 0.6102 - accuracy: 0.8068 - val_loss: 1.7839 - val_accuracy: 0.6114
Epoch 5/10
469/469 [==============================] - 21s 45ms/step - loss: 0.5670 - accuracy: 0.8263 - val_loss: 1.8409 - val_accuracy: 0.6156
Epoch 6/10
469/469 [==============================] - 20s 43ms/step - loss: 0.5412 - accuracy: 0.8346 - val_loss: 1.9607 - val_accuracy: 0.6187
Epoch 7/10
469/469 [==============================] - 20s 42ms/step - loss: 0.5019 - accuracy: 0.8468 - val_loss: 1.9545 - val_accuracy:
```