### Codio Activity 3.5: Sorting the result of `groupby`

**Expected Time**: 60 Minutes

**Total Points**: 10

This activity focuses on sorting the results of a groupby operation. The dataset is again the gapminder data from plotly.  Your tasks involve aggregating and sorting as well as plotting summary results using `matplotlib`, `seaborn`, and `plotly`.

#### Index:

- [Problem 1](#Problem-1:-Lowest-to-Highest-Average-Life-Expectancy)
- [Problem 2](#Problem-2:-Sorting-GDP)
- [Problem 3](#Problem-3:-Sorting-Results-of-Multiple-Aggregates)
- [Problem 4](#Problem-4:-Sorting-Custom-Outputs)
- [Problem 5](#Problem-5:-Applying-the-\`column_range\`-function)

```python
import plotly.express as px
import pandas as pd
import numpy as np
import seaborn as sns
import matplotlib.pyplot as plt
```

```python
gapminder = px.data.gapminder()
```

Output:
```
NameError                                 Traceback (most recent call last)
/var/folders/y8/q3tyf2ls6hdcr5jmw1zdr2f40000gp/T/ipykernel_84326/1138654827.py in <module>
----> 1 gapminder = px.data.gapminder()

NameError: name 'px' is not defined
```

```python
gapminder.head(2)
```

```python
gapminder.info()
```

[Back to top](#Index:) 

### Problem 1: Lowest to Highest Average Life Expectancy

**2 Points**

Using the gapminder data, group on the continent and find the average life expectancy by continent.  Sort the results from greatest to least using the `.sort_values` method.  Make sure to set the `ascending` argument to `False`.  Assign your results as a DataFrame to `ans1` below.  

<!-- grade_id: cell-9174b84554ea287c, points: 0, grade: false, solution: true -->
```python
### GRADED

ans1 = None

### BEGIN SOLUTION
#ans1 = gapminder.groupby('continent')[['lifeExp']].mean().sort_values(by = 'lifeExp', ascending = False)
ans1 = gapminder.groupby('continent')[['lifeExp']].agg('mean').sort_values('lifeExp', ascending=False)

### END SOLUTION

# Answer check
print(ans1)
print(type(ans1))
```

<!-- grade_id: cell-3667c338ac1d9a93, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
gapminder_ = px.data.gapminder()
ans1_ = gapminder_.groupby('continent')[['lifeExp']].mean().sort_values(by = 'lifeExp', ascending = False)
#
#
#
assert type(ans1_) == type(ans1)
pd.testing.assert_frame_equal(ans1, ans1_)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 2: Sorting GDP

**2 Points**

Which country has the highest average GDP?  Use groupby to answer this question and assign your answer as a one word string representing the name of the country to `ans2` below.

<!-- grade_id: cell-53f700d04cb1c94e, points: 0, grade: false, solution: true -->
```python
### GRADED

ans2 = None

### BEGIN SOLUTION
#ans2 = 'Kuwait' #gapminder.groupby('country')[['gdpPercap']].mean().sort_values(by = 'gdpPercap', ascending = False)
ans2 = gapminder.groupby('country', as_index=False)[['gdpPercap']].agg('mean').sort_values('gdpPercap', ascending=False).head(1).iloc[0,0]
### END SOLUTION

# Answer check
print(ans2)
print(type(ans2))
```

<!-- grade_id: cell-d18568269e3c1098, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
ans2_ = 'Kuwait'
#
#
#
assert type(ans2_) == type(ans2)
assert ans2_.lower() == ans2.lower()
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 3: Sorting Results of Multiple Aggregates

**2 Points**

For this problem, use the `['mean', 'std', 'median']` functions to determine these values for the Life Expectancy of each continent.  Save this DataFrame to `ans3a` below  To sort the resulting values by medians from greatest to least, first flatten and rename the columns using the code:

```python
ans3a.columns = [" ".join(c) for c in ans3a.columns.to_flat_index()]
```


Sort the results based on the median and assign your results as a DataFrame to `ans3b` below.  

<!-- grade_id: cell-52e240bec4e8a6a0, points: 0, grade: false, solution: true -->
```python
### GRADED

ans3a = None
ans3b = None

### BEGIN SOLUTION
#ans3a = gapminder.groupby('continent').agg(['mean', 'std', 'median'])[['lifeExp']]
#ans3a.columns = [" ".join(c) for c in ans3a.columns.to_flat_index()]
#ans3b = ans3a.sort_values(by = 'lifeExp median')
ans3a = gapminder.groupby('continent')[['lifeExp']].agg(['mean','std','median'])
ans3a.columns = [" ".join(c) for c in ans3a.columns.to_flat_index()]

ans3b = ans3a.sort_values('lifeExp median')
### END SOLUTION

# Answer check
print(ans3a)
print(ans3b)
```

<!-- grade_id: cell-e472f07102714ee4, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
gapminder_ = px.data.gapminder()
ans3a_ = gapminder_.groupby('continent').agg(['mean', 'std', 'median'])[['lifeExp']]
ans3a_.columns = [" ".join(c) for c in ans3a_.columns.to_flat_index()]
ans3b_ = ans3a_.sort_values(by = 'lifeExp median')
#
#
#
assert type(ans3a_) == type(ans3a)
pd.testing.assert_frame_equal(ans3a, ans3a_)
pd.testing.assert_frame_equal(ans3b, ans3b_)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 4: Sorting Custom Outputs

**2 Points**

Complete the function `column_range` that takes in a column, and returns the range (max - min) of a column in a DataFrame. 

<!-- grade_id: cell-ef54140f4c5e26da, points: 0, grade: false, solution: true -->
```python
### GRADED

def column_range():
    '''
    This function takes in a pandas series and returns
    the range of the series (max - min) as a float.
    '''
    pass

### BEGIN SOLUTION
#def column_range(col):
#    return col.max() - col.min()
def column_range(s):
    '''
    This function takes in a pandas series and returns
    the range of the series (max - min) as a float.
    '''
    offset=s.max() - s.min()
    return offset
### END SOLUTION

# Answer check
column_range(gapminder['lifeExp'])
```

<!-- grade_id: cell-1f39c6bb5e22442b, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
gapminder_ = px.data.gapminder()
def column_range_(col):
    return col.max() - col.min()


ans4 = column_range(gapminder['lifeExp'])
ans4_ = column_range_(gapminder['lifeExp'])
#
#
#
assert type(ans4) == type(ans4_)
assert ans4 == ans4_
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 5: Applying the `column_range` function

**2 Points**

Use the `column_range` function to first groupby continent and then apply the function to the GDP feature.  Sort the results from greatest to least.  Assign your results as a DataFrame to `ans5` below.

<!-- grade_id: cell-c44321bb8f95b1c6, points: 0, grade: false, solution: true -->
```python
### GRADED

ans5 = None

### BEGIN SOLUTION
#ans5 = gapminder.groupby('continent')[['gdpPercap']].apply(column_range).sort_values(by = 'gdpPercap', ascending = False)
ans5 = gapminder.groupby('continent')[['gdpPercap']].apply(column_range).sort_values(by = 'gdpPercap', ascending=False)
### END SOLUTION

# Answer check
ans5
```

<!-- grade_id: cell-c48d7e9f2e76e4fb, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
gapminder_ = px.data.gapminder()
def column_range_(col):
    return col.max() - col.min()

ans5_ = gapminder_.groupby('continent')[['gdpPercap']].apply(column_range).sort_values(by = 'gdpPercap', ascending = False)

#
#
#
pd.testing.assert_frame_equal(ans5_, ans5)
### END HIDDEN TESTS
```